# Python Final Exam
Epita international masters - Fall 2022 - Duquesne

* All the exercices are going to be graded automatically.
* Remember to fill in your name correctly.
* If you delete cells that are commented as #Do not delete the automatic grading will directly award no points for that exercise.
* Don't forget to remove raise NotImplementedError() if you code a solution for a function.
* Automatic plagiarism will also be run on your notebook.
* As in any academic context, you should cite your sources and collaborators.
* Only built-in python modules are allowed: https://docs.python.org/3/py-modindex.html

**Duration of the exam is 1 hour 30.**

**Access to the internet and class materials are allowed**

**Discussion between peers is prohibited**

```python
#Do not delete
name = "Bhava Sai Teja KANCHERLA"
sources = "GeekForGeeks"" ""PythonTutorialPoint"
collaborators = ""
```

## Exercise 1 (3 points)
A company has it's employees data formatted in a list of tuples. A tuple conists of the first name, the last name and the age of the person.
For example:
`[('Leo', 'Vernel', 38), ('Marjane', 'Nguyen', 23), ('Abigail', 'Deemle', 23)]`

The boss wants us to code a function `make_banner()` that will take the list of tuples as input and format a nice string so we can print it out and make a banner to congratulate our employees. The names should appear depending on the age, youngest to oldest and if we have people of the same age then their names should be sorted alphabetically. The final list of sorted names should be seperated with a `' - '` between them.

With our previous list of tuples tuple we would have as output string:

`Abigail - Marjane - Leo`

Abigail and Marjane are the youngest so they appear first. They are both 23 so we sort their names alphabetically. Leo is older so he appears after.

<!-- grade_id: cell-4b6c0ac771f428af, points: 0, grade: false, solution: true -->
```python
#Do not delete
def make_banner(employees: list) -> str:
    """
    Takes list of tuples of employee names and ages and formats a banner
    with names sorted depending on the age first and alphabetically 
    second, if ages are equivalent.
    
    #######
    input: employees (list)
    output: banner (str)
    #######
    """
    # YOUR CODE HERE
    
    emp_sorted = sorted(employees, key=lambda e: (e[2], e[0]))
    return ' - '.join([f"{emp[0]}" for emp in emp_sorted])
        
    
```

```python
employees = [('Leo', 'Vernel', 38), ('Marjane', 'Nguyen', 23), ('Abigail', 'Deemle', 23)]

print(make_banner(employees))
```

Output:
```
Abigail - Marjane - Leo
```

<!-- grade_id: cell-0af4e645eb1156c0, points: 3, grade: true, solution: false -->
```python
#Do not delete
employees = [('Leo', 'Vernel', 38), ('Marjane', 'Nguyen', 23), ('Abigail', 'Deemle', 23)]
assert make_banner(employees) == 'Abigail - Marjane - Leo'

employees = [('Kendry', 'Michel', 44), ('Lilian', 'Genery', 36), ('Kim', 'Lecosta', 21), ('Arabelle', 'Rimey', 52)]
assert make_banner(employees) == 'Kim - Lilian - Kendry - Arabelle'

employees = [('Kendry', 'Michel', 23), ('Lilian', 'Genery', 23), ('Kim', 'Lecosta', 23), ('Arabelle', 'Rimey', 23)]
assert make_banner(employees) == 'Arabelle - Kendry - Kim - Lilian'

employees = []
assert make_banner(employees) == ''

employees = [('Kendry', 'Michel', 44)]
assert make_banner(employees) == 'Kendry'

```

## Exercise 2 (3 points)
The [collatz sequence](https://simple.wikipedia.org/wiki/Collatz_sequence)
goes like this:

- Start by any number greater than zero.
- If your number is even, divide it by two.
- If your number is odd, multiply by three and add one.

For example, starting by `10` we have:

   ```10 → 5 → 16 → 8 → 4 → 2 → 1```

Write a function, named `collatz_length`, that given a number, returns
the length of the sequence before reaching 1, for example:

```python
>>> collatz_length(10)
6
```

Because it takes 6 steps (6 `→` in the previous example) to reach 1.



<!-- grade_id: cell-4c907e986917720f, points: 0, grade: false, solution: true -->
```python
#Do not delete
def collatz_length(n: int) -> int:
    """
    Calculates the numbers of steps in the collatz sequence
    
    #######
    input: n (integer)
    output: steps (integer)
    #######
    """
    # YOUR CODE HERE
    lgth = 0
    while n != 1:
        if n % 2 == 0:
            n = n // 2
        else:
            n = 3 * n + 1
        lgth += 1
    return lgth
```

```python
print(collatz_length(31328))
```

Output:
```
54
```

<!-- grade_id: cell-8cc20a4b16913931, points: 3, grade: true, solution: false -->
```python
#Do not delete
assert collatz_length(10) == 6
assert collatz_length(5) == 5
assert collatz_length(0) == 0
assert collatz_length(2) == 1
assert collatz_length(86) == 30
assert collatz_length(-123) == 0
```

## Exercise 3 (3 points)

Write a function called `get_string_index` that returns the index of  the first string element inside an array of elements of different types. The string can hold any value.

If there is no string element in the array, return `-1`. 

For example in the array `[1,2,'hello']` the string is at index `2`

<!-- grade_id: cell-fc55a4974b2473be, points: 0, grade: false, solution: true -->
```python
#Do not delete
def get_string_index(my_list: list) -> int:
    """
    Retruns the index of the first element of type string. 
    If no element are of type string inside the array returns -1
    #######
    input: my_list (list)
    output: index (integer)
    #######
    """
    # YOUR CODE HEREfor x, lst_itm in enumerate(my_list):
    for x, lst_itm in enumerate(my_list):
      if isinstance(lst_itm, str):
        return x
    return -1
```

```python
print(get_string_index([1,2,'hello']))
```

<!-- grade_id: cell-bf5435a01491ca62, points: 3, grade: true, solution: false -->
```python
#Do not delete
assert get_string_index([1,2,'hello']) == 2
assert get_string_index([1,2,'']) == 2
assert get_string_index([1,2,'? *,']) == 2
assert get_string_index([1,2,3,4,5,'a','d',1,2,3]) == 5
assert get_string_index(['e',True,3,4,5,'a','d',[2,3]]) == 0
assert get_string_index([False,True,3,4,5,6,2,[2,3]]) == -1
assert get_string_index([]) == -1
```

## Exercise 4 (6 points)

Write a simple function, named `is_anagram`, taking two strings, and returning a boolean value.

The function shall return `True` if the letters of one word are a rearrangement of the letters of the other, `False` otherwise.

Superfluous (or missing) spaces, quotes, dashes, ... are allowed: `funeral` is an anagram of `real fun`.

Capitalized letter are equivalent to lower case letters: `Madam Curie` is an anagram of `Radium came`.

Diactirics of any language are ignored: `crâné` is an anagram of `crane`.

*hint: using the module unicodedata might be helpful*

<!-- grade_id: cell-ecea43aafb0fd613, points: 0, grade: false, solution: true -->
```python
#Do not delete
import unicodedata
def is_anagram(word1: str, word2: str) -> bool:
    """
    Checks wheter two words are anagrams
    #######
    input: word1 (string), word2 (string)
    output: anagram (boolean)
    #######
    
    """
    # YOUR CODE HERE
    nfkd_val, ascii_val, ignore ="NFKD", "ascii", "ignore"
    word1 = unicodedata.normalize(nfkd_val, word1).encode(ascii_val, ignore).lower()
    word2 = unicodedata.normalize(nfkd_val, word2).encode(ascii_val, ignore).lower()
    return sorted(word1) == sorted(word2)
```

```python
print(is_anagram("evil", "vile"))
```

<!-- grade_id: cell-709a485b40a691c3, points: 1, grade: true, solution: false -->
```python
#Do not delete
assert is_anagram("evil", "vile") == True
assert is_anagram("ab", "aba") == False
assert is_anagram("abba", "baba") == True
assert is_anagram("hello", "bye") == False
```

<!-- grade_id: cell-4dfb98cc0db28211, points: 2, grade: true, solution: false -->
```python
#Do not delete
assert is_anagram("Madam Curie", "Radium came") == True
assert is_anagram("funeral", "real fun") == True
assert is_anagram("manoir", "Romina") == True
```

<!-- grade_id: cell-1e847e9ebcdd64f1, points: 3, grade: true, solution: false -->
```python
#Do not delete
assert is_anagram("New York Times", "monkeys write") == True
assert is_anagram("McDonald's restaurants", "Uncle Sam's standard rot") == True
assert is_anagram("crane", "crâné") == True
assert is_anagram("àbba", "Baba!") == True
```

## Exercise 5 (5 points)
In this exercise you should code a function `get_triple_number()` that takes an unordered list of integers. All integers are present once in the range 1 to n (n included), execpet one number that appears three times in the list. The function should return the value of that number appearing thrice.

For example in the array `[1,6,3,7,2,5,4,2,2]` all numbers between 1 and 7 included appear once in the array, except the number 2 that appears three times.

You can assume that the list will have a minimum size of three elements and only have numbers strictly bigger than 0.

The function should run in O(n) (at worst).

Having a function that works is worth 3 points and having a function that runs in O(n) is worth 2 points.

<!-- grade_id: cell-b4048dbe6b81c081, points: 0, grade: false, solution: true -->
```python
#Do not delete
def get_triple_number(numbers: list) -> int:
    """
    Takes an unordered list of all integers in range 1 to n. 
    Finds the only element that appears three times in the list.
    #######
    input: numbers (list)
    output: triple_number (int)
    #######
    
    """
    # YOUR CODE HERE
    unq_nums = set()
    for num in numbers:
        if num in unq_nums:
            unq_nums.remove(num)
            return num
        unq_nums.add(num)
```

```python
print(get_triple_number([1,6,3,7,2,5,4,2,2]))
```

<!-- grade_id: cell-00a4f63b5f4297d6, points: 3, grade: true, solution: false -->
```python
#Do not delete
assert get_triple_number([1,6,3,7,2,5,4,2,2]) == 2
assert get_triple_number([1,5,4,5,3,2,5]) == 5
assert get_triple_number([7,8,5,3,2,4,8,1,6,8]) == 8
assert get_triple_number([1,1,1]) == 1
assert get_triple_number([2,2,1,2]) == 2

```

<!-- grade_id: cell-1da33cb027d5d00e, points: 2, grade: true, solution: false -->
```python
#If this cells outputs an error or takes more than a second to execute 
#then it means your code is too slow...

#Do not delete

import random
import time
from statistics import median

class RunTimeTooSlow(Exception):
    "Raised when the code is too slow"
    pass

try:
    n = random.randint(1, 10001)
    l = [ i for i in range(1, 10001)] + [n] + [n]
    times = []
    for _ in range(25):
        start = time.time()
        get_triple_number(l)
        stop = time.time()
        random.shuffle(l)
        times.append(stop-start)
    print("code runs in {} seconds".format(median(times)))
    if median(times) > 0.01:
        raise RunTimeTooSlow
    else:
        print("Code OK ! Runs in O(n) ")
        
except RunTimeTooSlow:
    print("Exception occurred: RunTimeTooSlow")
```